<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">

    <name>
      <text>Shared: Complete a partially filled linear-relation table</text>
    </name>


    <questiontext format="html">
      <text>&lt;p&gt;
A linear relation is written as
\[
y-({@c@})={@m@}x.
\]
&lt;/p&gt;

&lt;p&gt;
Complete the three missing values in the table.
Some values have already been provided to help
you check the pattern.
&lt;/p&gt;

&lt;table border="1"
       cellpadding="8"
       style="border-collapse:collapse;
              text-align:center;"&gt;

&lt;tr&gt;
  &lt;th&gt;\(x\)&lt;/th&gt;
  &lt;td&gt;\({@x1@}\)&lt;/td&gt;
  &lt;td&gt;\({@x2@}\)&lt;/td&gt;
  &lt;td&gt;\({@x3@}\)&lt;/td&gt;
  &lt;td&gt;\({@x4@}\)&lt;/td&gt;
  &lt;td&gt;\({@x5@}\)&lt;/td&gt;
&lt;/tr&gt;

&lt;tr&gt;
  &lt;th&gt;\(y\)&lt;/th&gt;

  &lt;td&gt;
    \({@y1@}\)
  &lt;/td&gt;

  &lt;td&gt;
    [[input:ans1]]
    [[validation:ans1]]
    [[feedback:prt1]]
  &lt;/td&gt;

  &lt;td&gt;
    \({@y3@}\)
  &lt;/td&gt;

  &lt;td&gt;
    [[input:ans2]]
    [[validation:ans2]]
    [[feedback:prt2]]
  &lt;/td&gt;

  &lt;td&gt;
    [[input:ans3]]
    [[validation:ans3]]
    [[feedback:prt3]]
  &lt;/td&gt;

&lt;/tr&gt;

&lt;/table&gt;

&lt;p&gt;
Remember that the equation is not yet written
with \(y\) by itself.
&lt;/p&gt;
      </text>
    </questiontext>


    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
The given relation is
\[
y-({@c@})={@m@}x.
\]
&lt;/p&gt;

&lt;p&gt;
First isolate \(y\) by adding \({@c@}\)
to both sides:
\[
y={@m@}x+({@c@}).
\]
&lt;/p&gt;

&lt;p&gt;
For \(x={@x2@}\):
\[
y=({@m@})({@x2@})+({@c@})
={@y2@}.
\]
&lt;/p&gt;

&lt;p&gt;
For \(x={@x4@}\):
\[
y=({@m@})({@x4@})+({@c@})
={@y4@}.
\]
&lt;/p&gt;

&lt;p&gt;
For \(x={@x5@}\):
\[
y=({@m@})({@x5@})+({@c@})
={@y5@}.
\]
&lt;/p&gt;

&lt;p&gt;
Therefore the completed row is
\[
{@y1@},\;
{@y2@},\;
{@y3@},\;
{@y4@},\;
{@y5@}.
\]
&lt;/p&gt;
      </text>
    </generalfeedback>


    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />


    <stackversion>
      <text>2025012100</text>
    </stackversion>


    <questionvariables>
      <text>
/*
Shared Question Bank
Linear Relations
QLIN-02

This question is curriculum-independent.

Pedagogical distinction from QLIN-01:
- relation is not initially solved for y;
- table is partially completed;
- learner completes selected missing entries;
- existing values provide pattern evidence.
*/

m:rand_with_prohib(-4,4,[0]);

/*
Keep the constant non-zero so the targeted
sign misconception remains distinguishable
from the correct answer.
*/

c:rand_with_prohib(-5,5,[0]);


/*
Fixed x-values preserve a manageable
cognitive load while the learner focuses on
interpreting the alternative equation form.
*/

x1:-3;
x2:-1;
x3:0;
x4:2;
x5:4;


/*
Teacher answers.
*/

y1:m*x1+c;
y2:m*x2+c;
y3:m*x3+c;
y4:m*x4+c;
y5:m*x5+c;


/*
Common misconception:

From

    y-c = mx

the learner incorrectly subtracts c again
and uses

    y = mx-c

instead of adding c to both sides.
*/

wrong1:m*x2-c;
wrong2:m*x4-c;
wrong3:m*x5-c;
      </text>
    </questionvariables>


    <specificfeedback format="html">
      <text />
    </specificfeedback>


    <questionnote format="html">
      <text>&lt;p&gt;
Shared Question Bank |
Linear Relations |
QLIN-02
&lt;/p&gt;

&lt;p&gt;
m={@m@}; c={@c@}
&lt;/p&gt;

&lt;p&gt;
given=
({@x1@},{@y1@}),
({@x3@},{@y3@})
&lt;/p&gt;

&lt;p&gt;
missing_answers=
{@y2@},
{@y4@},
{@y5@}
&lt;/p&gt;
      </text>
    </questionnote>


    <questiondescription format="html">
      <text>&lt;p id="bank"&gt;
Shared Curriculum-Independent Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
QLIN-02
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
linear-relations
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
tables-of-values
&lt;/p&gt;

&lt;p id="role"&gt;
foundation
&lt;/p&gt;

&lt;p id="interaction"&gt;
partial-table-completion
&lt;/p&gt;

&lt;p id="description"&gt;
Learners interpret a linear relation written
in a form that is not solved for \(y\), then
complete selected missing values in a
partially populated table.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Original STACK implementation. Table-completion
interaction informed by openly accessible
Numbas interaction patterns. No Numbas
question wording or mathematical content is
copied into this item.
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained outside
the question source. This question may be
mapped independently to Ghana, Kenya, or
future curriculum profiles.
&lt;/p&gt;
      </text>
    </questiondescription>


    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>


    <prtcorrect format="html">
      <text>&lt;p&gt;Correct.&lt;/p&gt;</text>
    </prtcorrect>


    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;
Some of the missing values are correct.
Check how you rearranged the equation before
substituting each \(x\)-value.
&lt;/p&gt;</text>
    </prtpartiallycorrect>


    <prtincorrect format="html">
      <text>&lt;p&gt;
That value is not correct.
First isolate \(y\), then substitute the
corresponding value of \(x\).
&lt;/p&gt;</text>
    </prtincorrect>


    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />


    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>y2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>


    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>y4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>


    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>y5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>


    <prt>
      <name>prt1</name>
      <value>0.3333333</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>

      <node>
        <name>0</name>
        <description>
Checks the missing value for x2.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>y2</tans>
        <testoptions />
        <quiet>0</quiet>

        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks whether c was moved with the wrong sign.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong1</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
Check the sign of the constant.
From
\[
y-c=mx,
\]
add \(c\) to both sides:
\[
y=mx+c.
\]
Do not subtract \(c\) a second time.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>


    <prt>
      <name>prt2</name>
      <value>0.3333333</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>

      <node>
        <name>0</name>
        <description>
Checks the missing value for x4.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>y4</tans>
        <testoptions />
        <quiet>0</quiet>

        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks whether c was moved with the wrong sign.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrong2</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
You may have moved the constant with the
wrong sign. Since \(y-c=mx\), adding \(c\)
to both sides gives \(y=mx+c\).
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>


    <prt>
      <name>prt3</name>
      <value>0.3333334</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>

      <node>
        <name>0</name>
        <description>
Checks the missing value for x5.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>y5</tans>
        <testoptions />
        <quiet>0</quiet>

        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks whether c was moved with the wrong sign.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>wrong3</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
Rearrange the equation carefully.
The constant is added when \(y\) is isolated:
\[
y=mx+c.
\]
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
    </prt>


    <qtest>
      <testcase>1</testcase>

      <description>
Teacher answers should receive full credit
for all three missing table values.
      </description>

      <testinput>
        <name>ans1</name>
        <value>y2</value>
      </testinput>

      <testinput>
        <name>ans2</name>
        <value>y4</value>
      </testinput>

      <testinput>
        <name>ans3</name>
        <value>y5</value>
      </testinput>

      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>

      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>

      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>

    </qtest>

  </question>
</quiz>
