<quiz>
  <question type="stack">
    <name>
      <text>Shared: Read missing coordinates from two linear graphs</text>
    </name>

    <questiontext format="html">
      <text>&lt;p&gt;
Two linear relations,
&lt;strong&gt;Relation A&lt;/strong&gt; and
&lt;strong&gt;Relation B&lt;/strong&gt;,
are shown on the Cartesian plane.
&lt;/p&gt;

&lt;p&gt;
Use the graph to complete the two
ordered pairs.
&lt;/p&gt;

&lt;center&gt;

[[jsxgraph]]

JXG.Options.text.useMathJax = true;

var board = JXG.JSXGraph.initBoard(
    divid,
    {
        axis: true,
        showCopyright: false,
        boundingbox: [-6,6,6,-6],
        keepaspectratio: true,
        grid: true
    }
);


var lineA = board.create(
    'functiongraph',
    [
        function(x) {
            return {#m1#}*x + {#c1#};
        }
    ],
    {
        fixed: true,
        strokeWidth: 3
    }
);


var lineB = board.create(
    'functiongraph',
    [
        function(x) {
            return {#m2#}*x + {#c2#};
        }
    ],
    {
        fixed: true,
        strokeWidth: 3,
        dash: 2
    }
);


board.create(
    'point',
    [{#xA#},{#yA#}],
    {
        name: 'A',
        fixed: true,
        size: 4,
        showInfobox: false
    }
);


board.create(
    'point',
    [{#xB#},{#yB#}],
    {
        name: 'B',
        fixed: true,
        size: 4,
        showInfobox: false
    }
);


board.unsuspendUpdate();

[[/jsxgraph]]

&lt;/center&gt;


&lt;p&gt;
Point \(A\) lies on
&lt;strong&gt;Relation A&lt;/strong&gt;:
\[
A=({@xA@},\, ?)
\]

The missing \(y\)-coordinate is
[[input:ans1]]
[[validation:ans1]]
&lt;/p&gt;


&lt;p&gt;
Point \(B\) lies on
&lt;strong&gt;Relation B&lt;/strong&gt;:
\[
B=(?,\, {@yB@})
\]

The missing \(x\)-coordinate is
[[input:ans2]]
[[validation:ans2]]
&lt;/p&gt;

[[feedback:prt1]]
[[feedback:prt2]]
      </text>
    </questiontext>

    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
For point \(A\), read the vertical
coordinate of the marked point on
Relation A.
&lt;/p&gt;

&lt;p&gt;
Therefore,
\[
A=({@xA@},{@yA@}).
\]
&lt;/p&gt;

&lt;p&gt;
For point \(B\), read the horizontal
coordinate of the marked point on
Relation B.
&lt;/p&gt;

&lt;p&gt;
Therefore,
\[
B=({@xB@},{@yB@}).
\]
&lt;/p&gt;

&lt;p&gt;
Remember that an ordered pair is
written as
\[
(x,y).
\]
The horizontal coordinate comes first
and the vertical coordinate comes
second.
&lt;/p&gt;
      </text>
    </generalfeedback>

    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />

    <stackversion>
      <text>2023060500</text>
    </stackversion>

    <questionvariables>
      <text>/*
Shared Curriculum-Independent
Question Bank

Linear Relations
QLIN-07

Adaptation source:
NUMBAS-17290

Licence:
CC BY 4.0

The external source provides the
straight-line graph/equation question
family used as the adaptation basis.

The wording, variable generation,
two-relation structure, missing-pair
interaction, grading and feedback
below are implemented for this
shared STACK bank.
*/


/*
Generate two different straight
lines with small integer gradients
and intercepts.
*/

m1:rand_with_prohib(
    -2,
    2,
    [0]
);


m2:rand_with_prohib(
    -2,
    2,
    [0,m1]
);


c1:rand(5)-2;


c2:rand_with_prohib(
    -2,
    2,
    [c1]
);


/*
Choose readable integer points.

For A, x is shown and y is missing.

For B, y is shown and x is missing.
*/

xA:rand_with_prohib(
    -2,
    2,
    [0]
);


xB:rand_with_prohib(
    -2,
    2,
    [0,xA]
);


yA:m1*xA+c1;

yB:m2*xB+c2;


/*
Values generated by using the
wrong relation.

These support misconception
feedback.
*/

wrongA:m2*xA+c2;

wrongB:(yB-c1)/m1;


/*
Coordinate-order confusion values.
*/

confuseA:xA;

confuseB:yB;
      </text>
    </questionvariables>

    <specificfeedback format="html">
      <text>
[[feedback:prt1]]
[[feedback:prt2]]
      </text>
    </specificfeedback>

    <questionnote format="moodle_auto_format">
      <text>&lt;p&gt;
Shared Question Bank |
Linear Relations |
QLIN-07
&lt;/p&gt;

&lt;p&gt;
relationA:
m=\({@m1@}\);
c=\({@c1@}\)
&lt;/p&gt;

&lt;p&gt;
relationB:
m=\({@m2@}\);
c=\({@c2@}\)
&lt;/p&gt;

&lt;p&gt;
pointA:
({@xA@},{@yA@});
missing=y;
answer={@yA@}
&lt;/p&gt;

&lt;p&gt;
pointB:
({@xB@},{@yB@});
missing=x;
answer={@xB@}
&lt;/p&gt;

&lt;p&gt;
wrong_relation_A={@wrongA@};
wrong_relation_B={@wrongB@}
&lt;/p&gt;
      </text>
    </questionnote>

    <questiondescription format="moodle_auto_format">
      <text>&lt;p id="bank"&gt;
Shared Curriculum-Independent
Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
QLIN-07
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
linear-relations
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
coordinate-graphs
&lt;/p&gt;

&lt;p id="role"&gt;
conceptual
&lt;/p&gt;

&lt;p id="interaction"&gt;
missing-ordered-pair-from-graph
&lt;/p&gt;

&lt;p id="description"&gt;
Learners read two fixed linear graphs
and complete missing coordinates in
ordered pairs.
&lt;/p&gt;

&lt;p id="assessment-distinction"&gt;
Unlike QLIN-05 and QLIN-06, learners
do not construct the lines. They must
interpret already-drawn graph
information and distinguish horizontal
and vertical coordinates.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Adapted from the straight-line graph
question family represented by
Numbas question NUMBAS-17290,
"Graphs I: Linear - Line equation",
licensed CC BY 4.0.

The STACK implementation uses new
wording, generated values, a
two-relation missing-coordinate
interaction and diagnostic feedback.
&lt;/p&gt;

&lt;p id="source-url"&gt;
https://numbas.mathcentre.ac.uk/question/17290/
&lt;/p&gt;

&lt;p id="license"&gt;
CC BY 4.0
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained
outside the shared question source.
&lt;/p&gt;
      </text>
    </questiondescription>

    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>

    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>

    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>

    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>

    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />

    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>yA</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>

    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>xB</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>

    <prt>
      <name>prt1</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks the missing y-coordinate for A.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>yA</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Correct. You read the vertical
coordinate of point A correctly.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of Relation B for point A.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>wrongA</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
That value matches Relation B at the
given \(x\)-coordinate.

Point A belongs to Relation A. Read
the vertical coordinate from the
Relation A graph.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>2</name>
        <description>
Detects x/y coordinate-order confusion.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>confuseA</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
You appear to have used the
\(x\)-coordinate.

Ordered pairs are written as
\((x,y)\). The missing value for
A is the vertical \(y\)-coordinate.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Check point A on Relation A and read
its vertical coordinate carefully.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>

    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks the missing x-coordinate for B.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>xB</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Correct. You read the horizontal
coordinate of point B correctly.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of Relation A for point B.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>wrongB</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
That value corresponds to Relation A
at the shown \(y\)-coordinate.

Point B belongs to Relation B. Read
the horizontal coordinate from the
Relation B graph.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>2</name>
        <description>
Detects x/y coordinate-order confusion.
        </description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>confuseB</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
You appear to have used the
\(y\)-coordinate.

Ordered pairs are written as
\((x,y)\). The missing value for
B is the horizontal \(x\)-coordinate.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Check point B on Relation B and read
its horizontal coordinate carefully.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>

    <qtest>
      <testcase>1</testcase>
      <description>
Both missing coordinates should
receive full credit.
      </description>

      <testinput>
        <name>ans1</name>
        <value>yA</value>
      </testinput>

      <testinput>
        <name>ans2</name>
        <value>xB</value>
      </testinput>

      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>
prt1-1-T
        </expectedanswernote>
      </expected>

      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>
prt2-1-T
        </expectedanswernote>
      </expected>
    </qtest>
  <deployedseed>2029929245</deployedseed><deployedseed>905183159</deployedseed><deployedseed>1744478756</deployedseed><deployedseed>891006520</deployedseed><deployedseed>1594206364</deployedseed></question>
</quiz>