<quiz>
  <question type="stack">
    <name>
      <text>Shared: Table to graph simultaneous-equations mastery</text>
    </name>
    <questiontext format="html">
      <text>

&lt;p&gt;
This activity brings together the
main ideas from the linear-relations
progression.
&lt;/p&gt;

&lt;p&gt;
Consider the two relations
&lt;/p&gt;

\[
\text{Relation A: }
y={@m1@}x+({@c1@})
\]

\[
\text{Relation B: }
y={@m2@}x+({@c2@})
\]

&lt;p&gt;
First complete the missing values in
the two tables.
&lt;/p&gt;

&lt;p&gt;
&lt;strong&gt;Relation A&lt;/strong&gt;
&lt;/p&gt;

&lt;table border="1"
cellpadding="8"&gt;

&lt;tr&gt;
&lt;th&gt;x&lt;/th&gt;
&lt;th&gt;-2&lt;/th&gt;
&lt;th&gt;0&lt;/th&gt;
&lt;th&gt;2&lt;/th&gt;
&lt;/tr&gt;

&lt;tr&gt;
&lt;th&gt;y&lt;/th&gt;

&lt;td&gt;
[[input:ans1]]
[[validation:ans1]]
[[feedback:prt1]]
&lt;/td&gt;

&lt;td&gt;
{@a2@}
&lt;/td&gt;

&lt;td&gt;
[[input:ans2]]
[[validation:ans2]]
[[feedback:prt2]]
&lt;/td&gt;

&lt;/tr&gt;
&lt;/table&gt;


&lt;p&gt;
&lt;strong&gt;Relation B&lt;/strong&gt;
&lt;/p&gt;

&lt;table border="1"
cellpadding="8"&gt;

&lt;tr&gt;
&lt;th&gt;x&lt;/th&gt;
&lt;th&gt;-2&lt;/th&gt;
&lt;th&gt;0&lt;/th&gt;
&lt;th&gt;2&lt;/th&gt;
&lt;/tr&gt;

&lt;tr&gt;
&lt;th&gt;y&lt;/th&gt;

&lt;td&gt;
[[input:ans3]]
[[validation:ans3]]
[[feedback:prt3]]
&lt;/td&gt;

&lt;td&gt;
{@b2@}
&lt;/td&gt;

&lt;td&gt;
[[input:ans4]]
[[validation:ans4]]
[[feedback:prt4]]
&lt;/td&gt;

&lt;/tr&gt;
&lt;/table&gt;


&lt;p&gt;
Now construct both relations on the
same Cartesian plane.

Use \(A_1,A_2\) for Relation A and
\(B_1,B_2\) for Relation B.
&lt;/p&gt;

&lt;center&gt;

[[jsxgraph
input-ref-ans5="sA1"
input-ref-ans6="sA2"
input-ref-ans7="sB1"
input-ref-ans8="sB2"]]

JXG.Options.text.useMathJax = true;

var board = JXG.JSXGraph.initBoard(
    divid,
    {
        axis: true,
        showCopyright: false,
        boundingbox: [-6,6,6,-6],
        keepaspectratio: true,
        grid: true
    }
);


var a1 = board.create(
    'point',
    [
        {#start_ax1#},
        {#start_ay1#}
    ],
    {
        name:'A1',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var a2 = board.create(
    'point',
    [
        {#start_ax2#},
        {#start_ay2#}
    ],
    {
        name:'A2',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var lineA = board.create(
    'line',
    [
        a1,
        a2
    ],
    {
        fixed:true,
        straightFirst:true,
        straightLast:true,
        strokeWidth:2
    }
);


var b1 = board.create(
    'point',
    [
        {#start_bx1#},
        {#start_by1#}
    ],
    {
        name:'B1',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var b2 = board.create(
    'point',
    [
        {#start_bx2#},
        {#start_by2#}
    ],
    {
        name:'B2',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var lineB = board.create(
    'line',
    [
        b1,
        b2
    ],
    {
        fixed:true,
        straightFirst:true,
        straightLast:true,
        strokeWidth:2,
        dash:2
    }
);


stack_jxg.bind_point(
    sA1,
    a1
);

stack_jxg.bind_point(
    sA2,
    a2
);

stack_jxg.bind_point(
    sB1,
    b1
);

stack_jxg.bind_point(
    sB2,
    b2
);


board.unsuspendUpdate();

[[/jsxgraph]]

&lt;/center&gt;


&lt;p style="visibility:hidden"&gt;
[[input:ans5]]
[[validation:ans5]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans6]]
[[validation:ans6]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans7]]
[[validation:ans7]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans8]]
[[validation:ans8]]
&lt;/p&gt;


[[feedback:prt5]]

[[feedback:prt6]]




&lt;p&gt;
Identify the point where the two
relations intersect.

Enter the ordered pair as
\[
[x,y].
\]
&lt;/p&gt;

&lt;p&gt;
[[input:ans9]]
[[validation:ans9]]
[[feedback:prt7]]
&lt;/p&gt;


&lt;p&gt;
Finally, state the solution of the
simultaneous equations as the ordered
pair \([x,y]\).
&lt;/p&gt;

&lt;p&gt;
[[input:ans10]]
[[validation:ans10]]
[[feedback:prt8]]
&lt;/p&gt;

      </text>
    </questiontext>
    <generalfeedback format="html">
      <text>
&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
Equation A is
\[
y={@m1@}x+({@c1@}).
\]
&lt;/p&gt;

&lt;p&gt;
Equation B is
\[
y={@m2@}x+({@c2@}).
\]
&lt;/p&gt;

&lt;p&gt;
After plotting both equations,
their graphs meet at
\[
({@xI@},{@yI@}).
\]
&lt;/p&gt;

&lt;p&gt;
Therefore the graphical solution of
the simultaneous equations is
\[
[{@xI@},{@yI@}].
\]
&lt;/p&gt;
</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>
/*
Shared Curriculum-Independent
Question Bank

Linear Relations
QLIN-12

Final integrated mastery activity.

Verified open-source bases:

NUMBAS-17290
CC BY 4.0
Straight-line graph construction.

NUMBAS-17144
CC BY 4.0
Simultaneous-equation structure.

NUMBAS-40889
CC BY 4.0
Simultaneous-equation solution
reference.
*/


/*
Choose a small intersection first.

x is non-zero.

y differs from x so a reversed
ordered pair remains detectable.
*/

xI:rand_with_prohib(
    -1,
    1,
    [0]
);

yI:rand_with_prohib(
    -2,
    2,
    [xI]
);


/*
Generate two different,
non-horizontal gradients.
*/

m1:rand_with_prohib(
    -2,
    2,
    [0]
);

m2:rand_with_prohib(
    -2,
    2,
    [0,m1]
);


/*
Construct both intercepts so the
relations meet at the selected
intersection.
*/

c1:yI-m1*xI;

c2:yI-m2*xI;


/*
Equations supplied to the learner.
*/

equation1:y=m1*x+c1;

equation2:y=m2*x+c2;


/*
Table x-values.
*/

x1:-2;

x2:0;

x3:2;


/*
Relation A table.
*/

a1:m1*x1+c1;

a2:m1*x2+c1;

a3:m1*x3+c1;


/*
Relation B table.
*/

b1:m2*x1+c2;

b2:m2*x2+c2;

b3:m2*x3+c2;


/*
Wrong-relation diagnostics.

At each missing location, use the
other equation.
*/

wrongA1:m2*x1+c2;

wrongA3:m2*x3+c2;

wrongB1:m1*x1+c1;

wrongB3:m1*x3+c1;


/*
Teacher graph points.

Any two distinct points on the
correct relation are accepted by
the graph PRTs.
*/

ta1:[0,c1];

ta2:[1,m1+c1];

tb1:[0,c2];

tb2:[1,m2+c2];


/*
Intersection and final solution.
*/

intersection:[xI,yI];

reversed:[yI,xI];

solution:[xI,yI];


/*
Initial draggable locations.
*/

start_ax1:-5;

start_ay1:5;

start_ax2:-4;

start_ay2:4;

start_bx1:4;

start_by1:4;

start_bx2:5;

start_by2:5;
      </text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt5]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>
&lt;p&gt;
Shared Question Bank |
Linear Relations |
QLIN-12
&lt;/p&gt;

&lt;p&gt;
Equation A:
m=\({@m1@}\);
c=\({@c1@}\)
&lt;/p&gt;

&lt;p&gt;
Equation B:
m=\({@m2@}\);
c=\({@c2@}\)
&lt;/p&gt;

&lt;p&gt;
teacherA:
{@ta1@},
{@ta2@}
&lt;/p&gt;

&lt;p&gt;
teacherB:
{@tb1@},
{@tb2@}
&lt;/p&gt;

&lt;p&gt;
intersection=
\({@intersection@}\)
&lt;/p&gt;

&lt;p&gt;
reversed=
\({@reversed@}\)
&lt;/p&gt;
</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text>
&lt;p id="bank"&gt;
Shared Curriculum-Independent
Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
QLIN-12
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
linear-relations
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
coordinate-graphs
&lt;/p&gt;

&lt;p id="tertiary-concept"&gt;
simultaneous-equations
&lt;/p&gt;

&lt;p id="role"&gt;
mastery_check
&lt;/p&gt;

&lt;p id="interaction"&gt;
table-graph-intersection-mastery
&lt;/p&gt;

&lt;p id="description"&gt;
Learners complete missing table
values for two linear relations,
construct both graphs, identify their
intersection, and state the final
simultaneous-equation solution.
&lt;/p&gt;

&lt;p id="assessment-distinction"&gt;
QLIN-12 integrates the full shared
linear-relations progression into
one final mastery activity:
table completion, graph construction,
intersection identification, and
simultaneous-equation solution.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Integrated adaptation drawing on
verified CC BY 4.0 Numbas sources:
NUMBAS-17290 for straight-line graph
construction, NUMBAS-17144 for
simultaneous-equation structure, and
NUMBAS-40889 for simultaneous-equation
solution reference. The integrated
mastery workflow and diagnostic
structure are newly assembled for the
shared question bank.
&lt;/p&gt;

&lt;p id="license"&gt;
CC BY 4.0
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained
outside the shared question source.
&lt;/p&gt;
      </text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>a1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>a3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>b1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>b3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>tb1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>tb2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>intersection</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x,y]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>solution</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>[x,y]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>0.0750000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks a missing value in
Relation A.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>a1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt1-1-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt1-1-F
        </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of the other relation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrongA1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt1-2-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You appear to have used the other
linear relation.

Use Relation A for this
table entry.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt1-2-F
        </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Substitute the given \(x\)-value
carefully into Relation A.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.0750000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks a missing value in
Relation A.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>a3</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt2-1-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt2-1-F
        </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of the other relation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrongA3</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt2-2-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You appear to have used the other
linear relation.

Use Relation A for this
table entry.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt2-2-F
        </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Substitute the given \(x\)-value
carefully into Relation A.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.0750000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks a missing value in
Relation B.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>b1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt3-1-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt3-1-F
        </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of the other relation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>wrongB1</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt3-2-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You appear to have used the other
linear relation.

Use Relation B for this
table entry.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt3-2-F
        </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Substitute the given \(x\)-value
carefully into Relation B.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>0.0750000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks a missing value in
Relation B.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>b3</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt4-1-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt4-1-F
        </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects use of the other relation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>wrongB3</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt4-2-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You appear to have used the other
linear relation.

Use Relation B for this
table entry.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt4-2-F
        </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Substitute the given \(x\)-value
carefully into Relation B.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>0.2000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>
/*
Distance between the two A points.
A single duplicated point must not
define a line.
*/

distA:
    (ans5[1]-ans6[1])^2
    +
    (ans5[2]-ans6[2])^2;


/*
Residuals against Relation A.
Both draggable points must lie on
the required line.
*/

resA1:
    ans5[2]
    -
    (m1*ans5[1]+c1);

resA2:
    ans6[2]
    -
    (m1*ans6[1]+c1);

errorA:
    resA1^2
    +
    resA2^2;


/*
Swapped-line diagnostic:
A points were placed on Relation B.
*/

swapA1:
    ans5[2]
    -
    (m2*ans5[1]+c2);

swapA2:
    ans6[2]
    -
    (m2*ans6[1]+c2);

swapErrorA:
    swapA1^2
    +
    swapA2^2;
    </text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>
Checks that A1 and A2 are distinct.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>distA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt5-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Use two different points for
Relation A.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt5-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Checks Relation A.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>errorA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt5-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Relation A is correct.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>
prt5-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>
Checks whether Relation B was plotted
using the A points.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>swapErrorA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt5-3-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Your A-points appear to lie on
Equation B.

Equation A is
\[
y={@m1@}x+({@c1@}).
\]

Use \(A_1,A_2\) to graph
Equation A.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt5-3-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Equation A is not correct yet.

Check its gradient and
\(y\)-intercept, then place two
different points on that line.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>0.2000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>
/*
Distance between the two B points.
*/

distB:
    (ans7[1]-ans8[1])^2
    +
    (ans7[2]-ans8[2])^2;


/*
Residuals against Relation B.
*/

resB1:
    ans7[2]
    -
    (m2*ans7[1]+c2);

resB2:
    ans8[2]
    -
    (m2*ans8[1]+c2);

errorB:
    resB1^2
    +
    resB2^2;


/*
Swapped-line diagnostic.
*/

swapB1:
    ans7[2]
    -
    (m1*ans7[1]+c1);

swapB2:
    ans8[2]
    -
    (m1*ans8[1]+c1);

swapErrorB:
    swapB1^2
    +
    swapB2^2;
    </text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>
Checks that B1 and B2 are distinct.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>distB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt6-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Use two different points for
Relation B.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt6-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Checks Relation B.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>errorB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt6-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Relation B is correct.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>
prt6-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>
Checks whether Relation A was plotted
using the B points.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>swapErrorB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt6-3-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Your B-points appear to lie on
Equation A.

Equation B is
\[
y={@m2@}x+({@c2@}).
\]

Use \(B_1,B_2\) to graph
Equation B.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt6-3-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Equation B is not correct yet.

Check its gradient and
\(y\)-intercept, then place two
different points on that line.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>0.1500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <node>
        <name>0</name>
        <description>
Checks the graphical solution.
    </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>intersection</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt7-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. That ordered pair is the
common solution of both equations.
&lt;/p&gt;
      </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt7-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Detects reversed intersection
coordinates.
    </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>reversed</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt7-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
The coordinates appear to be reversed.

Remember that the graphical solution
is written in the order \([x,y]\).
&lt;/p&gt;
      </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt7-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Check where the two lines cross.

The point of intersection must lie on
both graphs.
&lt;/p&gt;
      </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>0.1500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks the final simultaneous-equation
solution.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>solution</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt8-1-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. This ordered pair satisfies
both linear equations.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt8-1-F
        </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Detects reversed final solution.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>reversed</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt8-2-T
        </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
The coordinates appear to be
reversed.

State the simultaneous-equation
solution in the order \([x,y]\).
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt8-2-F
        </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
The solution must satisfy both
relations simultaneously.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>
Complete integrated mastery
response.
      </description>
      <testinput>
        <name>ans1</name>
        <value>a1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>a3</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>b1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>b3</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>tb1</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>tb2</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>intersection</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>solution</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
  <deployedseed>424020968</deployedseed><deployedseed>154441093</deployedseed><deployedseed>1388385398</deployedseed><deployedseed>1271948827</deployedseed><deployedseed>2059664801</deployedseed></question>
</quiz>