<quiz>
  <question type="stack">
    <name>
      <text>Shared: Model and solve two plans graphically</text>
    </name>
    <questiontext format="html">
      <text>

&lt;p&gt;
Two service plans charge according
to the number of service units used.
&lt;/p&gt;

&lt;p&gt;
&lt;strong&gt;Plan A&lt;/strong&gt;
has a fixed fee of \({@bA@}\)
credits and charges \({@mA@}\)
credits per service unit.
&lt;/p&gt;

&lt;p&gt;
&lt;strong&gt;Plan B&lt;/strong&gt;
has a fixed fee of \({@bB@}\)
credits and charges \({@mB@}\)
credit per service unit.
&lt;/p&gt;

&lt;p&gt;
Let \(x\) represent the number of
service units and \(y\) represent
the total cost.
&lt;/p&gt;

&lt;p&gt;
First write the equation for each
plan.

Enter only the expression on the
right-hand side of \(y=\).
&lt;/p&gt;

&lt;p&gt;
&lt;strong&gt;Plan A:&lt;/strong&gt;

\(y=\)

[[input:ans1]]
[[validation:ans1]]

[[feedback:prt1]]
&lt;/p&gt;

&lt;p&gt;
&lt;strong&gt;Plan B:&lt;/strong&gt;

\(y=\)

[[input:ans2]]
[[validation:ans2]]

[[feedback:prt2]]
&lt;/p&gt;

&lt;p&gt;
Now use \(A_1,A_2\) to graph
Plan A and \(B_1,B_2\) to graph
Plan B.
&lt;/p&gt;

&lt;center&gt;

[[jsxgraph
input-ref-ans3="sA1"
input-ref-ans4="sA2"
input-ref-ans5="sB1"
input-ref-ans6="sB2"]]

JXG.Options.text.useMathJax = true;

var board = JXG.JSXGraph.initBoard(
    divid,
    {
        axis: true,
        showCopyright: false,
        boundingbox: [-1,15,6,-1],
        keepaspectratio: true,
        grid: true
    }
);


var a1 = board.create(
    'point',
    [
        {#start_ax1#},
        {#start_ay1#}
    ],
    {
        name:'A1',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var a2 = board.create(
    'point',
    [
        {#start_ax2#},
        {#start_ay2#}
    ],
    {
        name:'A2',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var lineA = board.create(
    'line',
    [
        a1,
        a2
    ],
    {
        fixed:true,
        straightFirst:true,
        straightLast:true,
        strokeWidth:2
    }
);


var b1 = board.create(
    'point',
    [
        {#start_bx1#},
        {#start_by1#}
    ],
    {
        name:'B1',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var b2 = board.create(
    'point',
    [
        {#start_bx2#},
        {#start_by2#}
    ],
    {
        name:'B2',
        snapToGrid:true,
        size:3,
        snapSizeX:0.1,
        snapSizeY:0.1,
        showInfobox:false
    }
);


var lineB = board.create(
    'line',
    [
        b1,
        b2
    ],
    {
        fixed:true,
        straightFirst:true,
        straightLast:true,
        strokeWidth:2,
        dash:2
    }
);


stack_jxg.bind_point(
    sA1,
    a1
);

stack_jxg.bind_point(
    sA2,
    a2
);

stack_jxg.bind_point(
    sB1,
    b1
);

stack_jxg.bind_point(
    sB2,
    b2
);


board.unsuspendUpdate();

[[/jsxgraph]]

&lt;/center&gt;


&lt;p style="visibility:hidden"&gt;
[[input:ans3]]
[[validation:ans3]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans4]]
[[validation:ans4]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans5]]
[[validation:ans5]]
&lt;/p&gt;

&lt;p style="visibility:hidden"&gt;
[[input:ans6]]
[[validation:ans6]]
&lt;/p&gt;


[[feedback:prt3]]

[[feedback:prt4]]




&lt;p&gt;
After constructing the two graphs,
enter their intersection as
\[
[x,y].
\]
&lt;/p&gt;

&lt;p&gt;
[[input:ans7]]
[[validation:ans7]]

[[feedback:prt5]]
&lt;/p&gt;

&lt;p&gt;
Now interpret the intersection.
&lt;/p&gt;

&lt;p&gt;
Number of service units:

[[input:ans8]]
[[validation:ans8]]

[[feedback:prt6]]
&lt;/p&gt;

&lt;p&gt;
Common total cost:

[[input:ans9]]
[[validation:ans9]]

[[feedback:prt7]]
&lt;/p&gt;

      </text>
    </questiontext>
    <generalfeedback format="html">
      <text>
&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
Equation A is
\[
y={@mA@}x+({@bA@}).
\]
&lt;/p&gt;

&lt;p&gt;
Equation B is
\[
y={@mB@}x+({@bB@}).
\]
&lt;/p&gt;

&lt;p&gt;
After plotting both equations,
their graphs meet at
\[
({@xI@},{@yI@}).
\]
&lt;/p&gt;

&lt;p&gt;
Therefore the graphical solution of
the simultaneous equations is
\[
[{@xI@},{@yI@}].
\]
&lt;/p&gt;
</text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>
/*
Shared Curriculum-Independent
Question Bank

Linear Relations
QLIN-11

Licensed adaptation source:
NUMBAS-17144

Licence:
CC BY 4.0

This task extends a licensed
simultaneous-equations question
family with independently written
contextual modelling, graphing,
solution and interpretation.
*/


xI:rand(3)+2;

mA:rand(2)+2;

mB:1;

bA:rand(2)+1;

bB:bA+(mA-mB)*xI;

yI:mA*xI+bA;


rhsA:mA*x+bA;

rhsB:mB*x+bB;


noFeeA:mA*x;

noFeeB:mB*x;


ta1:[0,bA];

ta2:[1,mA+bA];

tb1:[0,bB];

tb2:[1,mB+bB];


intersection:[xI,yI];

reversed:[yI,xI];


units:xI;

cost:yI;

reversed_units:yI;

reversed_cost:xI;


start_ax1:0;

start_ay1:13;

start_ax2:1;

start_ay2:12;

start_bx1:4;

start_by1:2;

start_bx2:5;

start_by2:3;
      </text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt3]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>
&lt;p&gt;
Shared Question Bank |
Linear Relations |
QLIN-11
&lt;/p&gt;

&lt;p&gt;
Equation A:
m=\({@mA@}\);
c=\({@bA@}\)
&lt;/p&gt;

&lt;p&gt;
Equation B:
m=\({@mB@}\);
c=\({@bB@}\)
&lt;/p&gt;

&lt;p&gt;
teacherA:
{@ta1@},
{@ta2@}
&lt;/p&gt;

&lt;p&gt;
teacherB:
{@tb1@},
{@tb2@}
&lt;/p&gt;

&lt;p&gt;
intersection=
\({@intersection@}\)
&lt;/p&gt;

&lt;p&gt;
reversed=
\({@reversed@}\)
&lt;/p&gt;
</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text>
&lt;p id="bank"&gt;
Shared Curriculum-Independent
Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
QLIN-11
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
linear-relations
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
coordinate-graphs
&lt;/p&gt;

&lt;p id="tertiary-concept"&gt;
simultaneous-equations
&lt;/p&gt;

&lt;p id="role"&gt;
challenge
&lt;/p&gt;

&lt;p id="interaction"&gt;
context-to-graph-simultaneous-equations
&lt;/p&gt;

&lt;p id="description"&gt;
Learners translate two contextual
linear relationships into equations,
construct both graphs, identify their
intersection, and interpret both
coordinates in context.
&lt;/p&gt;

&lt;p id="assessment-distinction"&gt;
Unlike QLIN-10, the equations are
not supplied. Learners must model
the context before graphing and
interpreting the solution.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Adapted from the simultaneous
equations question family represented
by Numbas question NUMBAS-17144,
"Mock exam Simultaneous linear
equations (variables)", licensed
CC BY 4.0.
&lt;/p&gt;

&lt;p id="license"&gt;
CC BY 4.0
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained
outside the shared question source.
&lt;/p&gt;
      </text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:green;"&gt;&lt;i class="fa fa-check"&gt;&lt;/i&gt;&lt;/span&gt; Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:orange;"&gt;&lt;i class="fa fa-adjust"&gt;&lt;/i&gt;&lt;/span&gt; Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;span style="font-size: 1.5em; color:red;"&gt;&lt;i class="fa fa-times"&gt;&lt;/i&gt;&lt;/span&gt; Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>rhsA</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>expression in x</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>rhsB</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>expression in x</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>tb1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>tb2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options />
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>intersection</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x,y]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>units</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>cost</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>number</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>0.1000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks Plan A equation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>rhsA</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. This models Plan A.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks for swapped plan.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>rhsB</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You entered the equation for
Plan B.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>2</name>
        <description>
Checks for omitted fixed fee.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>noFeeA</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You used the rate but omitted
the fixed fee.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Use the form \(mx+b\).
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.1000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks Plan B equation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>rhsB</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. This models Plan B.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks for swapped plan.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>rhsA</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You entered the equation for
Plan A.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>2</name>
        <description>
Checks for omitted fixed fee.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>noFeeB</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You used the rate but omitted
the fixed fee.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Use the form \(mx+b\).
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.2000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>
/*
Distance between the two A points.
A single duplicated point must not
define a line.
*/

distA:
    (ans3[1]-ans4[1])^2
    +
    (ans3[2]-ans4[2])^2;


/*
Residuals against Relation A.
Both draggable points must lie on
the required line.
*/

resA1:
    ans3[2]
    -
    (mA*ans3[1]+bA);

resA2:
    ans4[2]
    -
    (mA*ans4[1]+bA);

errorA:
    resA1^2
    +
    resA2^2;


/*
Swapped-line diagnostic:
A points were placed on Relation B.
*/

swapA1:
    ans3[2]
    -
    (mB*ans3[1]+bB);

swapA2:
    ans4[2]
    -
    (mB*ans4[1]+bB);

swapErrorA:
    swapA1^2
    +
    swapA2^2;
    </text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>
Checks that A1 and A2 are distinct.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>distA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt3-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Use two different points for
Relation A.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt3-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Checks Relation A.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>errorA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt3-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Relation A is correct.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>
prt3-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>
Checks whether Relation B was plotted
using the A points.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>swapErrorA</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt3-3-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Your A-points appear to lie on
Equation B.

Equation A is
\[
y={@mA@}x+({@bA@}).
\]

Use \(A_1,A_2\) to graph
Equation A.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt3-3-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Equation A is not correct yet.

Check its gradient and
\(y\)-intercept, then place two
different points on that line.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>0.2000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>
/*
Distance between the two B points.
*/

distB:
    (ans5[1]-ans6[1])^2
    +
    (ans5[2]-ans6[2])^2;


/*
Residuals against Relation B.
*/

resB1:
    ans5[2]
    -
    (mB*ans5[1]+bB);

resB2:
    ans6[2]
    -
    (mB*ans6[1]+bB);

errorB:
    resB1^2
    +
    resB2^2;


/*
Swapped-line diagnostic.
*/

swapB1:
    ans5[2]
    -
    (mA*ans5[1]+bA);

swapB2:
    ans6[2]
    -
    (mA*ans6[1]+bA);

swapErrorB:
    swapB1^2
    +
    swapB2^2;
    </text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>
Checks that B1 and B2 are distinct.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>distB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt4-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Use two different points for
Relation B.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt4-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Checks Relation B.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>errorB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt4-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Relation B is correct.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>
prt4-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>
Checks whether Relation A was plotted
using the B points.
    </description>
        <answertest>NumAbsolute</answertest>
        <sans>swapErrorB</sans>
        <tans>0</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt4-3-T
    </trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
Your B-points appear to lie on
Equation A.

Equation B is
\[
y={@mB@}x+({@bB@}).
\]

Use \(B_1,B_2\) to graph
Equation B.
&lt;/p&gt;</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt4-3-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Equation B is not correct yet.

Check its gradient and
\(y\)-intercept, then place two
different points on that line.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>0.2000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <node>
        <name>0</name>
        <description>
Checks the graphical solution.
    </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>intersection</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt5-1-T
    </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. That ordered pair is the
common solution of both equations.
&lt;/p&gt;
      </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>
prt5-1-F
    </falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Detects reversed intersection
coordinates.
    </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>reversed</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>
prt5-2-T
    </trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
The coordinates appear to be reversed.

Remember that the graphical solution
is written in the order \([x,y]\).
&lt;/p&gt;
      </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>
prt5-2-F
    </falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Check where the two lines cross.

The point of intersection must lie on
both graphs.
&lt;/p&gt;
      </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>0.1000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks contextual interpretation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>units</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. The horizontal coordinate gives the number of service units.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks reversed interpretation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>reversed_units</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You used the vertical coordinate. Use the horizontal coordinate.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Check the coordinates of the
intersection.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>0.1000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <node>
        <name>0</name>
        <description>
Checks contextual interpretation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>cost</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
Correct. The vertical coordinate gives the common total cost.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>

      <node>
        <name>1</name>
        <description>
Checks reversed interpretation.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>reversed_cost</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-2-T</trueanswernote>
        <truefeedback format="html">
          <text>
&lt;p&gt;
You used the horizontal coordinate. Use the vertical coordinate.
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>
&lt;p&gt;
Check the coordinates of the
intersection.
&lt;/p&gt;
          </text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>
Correct modelling, graphs,
intersection and interpretation.
      </description>
      <testinput>
        <name>ans1</name>
        <value>rhsA</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>rhsB</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>tb1</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>tb2</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>intersection</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>units</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>cost</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
    </qtest>
  <deployedseed>1320940429</deployedseed><deployedseed>773830893</deployedseed><deployedseed>136164401</deployedseed><deployedseed>193225268</deployedseed><deployedseed>1399545814</deployedseed></question>
</quiz>