<quiz>
  <question type="stack">

    <name>
      <text>
Shared: Find the interval of values that round to a target
      </text>
    </name>


    <questiontext format="html">
      <text>&lt;p&gt;
A whole number rounds to
\[
{@rounded_value@}
\]
when rounded to the nearest
&lt;strong&gt;{@place_name@}&lt;/strong&gt;.
&lt;/p&gt;

&lt;p&gt;
What is the
&lt;strong&gt;smallest whole number&lt;/strong&gt;
that could have produced this rounded value?
&lt;/p&gt;

&lt;p&gt;
[[input:ans1]]
[[validation:ans1]]
&lt;/p&gt;

&lt;p&gt;
What is the
&lt;strong&gt;largest whole number&lt;/strong&gt;
that could have produced this rounded value?
&lt;/p&gt;

&lt;p&gt;
[[input:ans2]]
[[validation:ans2]]
&lt;/p&gt;

&lt;p&gt;
[[feedback:prt1]]
&lt;/p&gt;
      </text>
    </questiontext>


    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
The rounded result is
\[
{@rounded_value@}
\]
to the nearest
&lt;strong&gt;{@place_name@}&lt;/strong&gt;.
&lt;/p&gt;

&lt;p&gt;
The lower boundary is halfway between
this rounded value and the previous
multiple:
\[
{@lower_answer@}.
\]
&lt;/p&gt;

&lt;p&gt;
The upper halfway value would round to
the next multiple under the half-up rule,
so the largest whole number in this
interval is one less:
\[
{@upper_answer@}.
\]
&lt;/p&gt;

&lt;p&gt;
Therefore the possible whole numbers are
\[
\boxed{
{@lower_answer@}
\text{ through }
{@upper_answer@}
}.
\]
&lt;/p&gt;
      </text>
    </generalfeedback>


    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />


    <stackversion>
      <text>2025012100</text>
    </stackversion>


    <questionvariables>
      <text>
/*
Shared Question Bank
Place Value, Rounding and Standard Form
PVRS-11

Challenge task:
reverse rounding by identifying the full
whole-number interval that produces a
given rounded result.

Original STACK question.
No source wording or mathematical content
is copied.
*/


/*
Choose the rounding place.

1 = nearest ten
2 = nearest hundred
3 = nearest thousand
*/

place_choice:
    rand_with_prohib(1,3,[]);

place_value:
    if place_choice=1 then 10
    elseif place_choice=2 then 100
    else 1000;

place_name:
    if place_choice=1 then "ten"
    elseif place_choice=2 then "hundred"
    else "thousand";


/*
Choose a positive target multiple.
*/

target_index:
    rand_with_prohib(
        120,
        850,
        []
    );

rounded_value:
    target_index
    * place_value;


/*
Under school half-up rounding:

smallest included whole number:
R - p/2

largest included whole number:
R + p/2 - 1

The exact upper halfway point belongs to
the next rounded value.
*/

lower_answer:
    rounded_value
    - place_value/2;

upper_answer:
    rounded_value
    + place_value/2
    - 1;


/*
Primary misconception:
assume the rounded result corresponds to
only one original value.
*/

wrong_unique_lower:
    rounded_value;

wrong_unique_upper:
    rounded_value;


/*
Secondary misconception:
include the upper halfway boundary even
though that value rounds upward to the
next multiple.
*/

wrong_upper_inclusive:
    rounded_value
    + place_value/2;
      </text>
    </questionvariables>


    <specificfeedback format="html">
      <text />
    </specificfeedback>


    <questionnote format="html">
      <text>&lt;p&gt;
Shared Question Bank |
Place Value and Rounding |
PVRS-11
&lt;/p&gt;

&lt;p&gt;
rounded_value={@rounded_value@}
&lt;/p&gt;

&lt;p&gt;
place_value={@place_value@}
&lt;/p&gt;

&lt;p&gt;
lower_answer={@lower_answer@}
&lt;/p&gt;

&lt;p&gt;
upper_answer={@upper_answer@}
&lt;/p&gt;
      </text>
    </questionnote>


    <questiondescription format="html">
      <text>&lt;p id="bank"&gt;
Shared Curriculum-Independent Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
PVRS-11
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
number-sense
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
reverse-rounding
&lt;/p&gt;

&lt;p id="role"&gt;
challenge
&lt;/p&gt;

&lt;p id="interaction"&gt;
reverse-rounding-interval
&lt;/p&gt;

&lt;p id="description"&gt;
Learners identify the smallest and largest
whole numbers that round to a specified
target value.
&lt;/p&gt;

&lt;p id="misconception"&gt;
Targets the belief that a rounded value
corresponds to only one original number,
and checks understanding of the upper
halfway boundary.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Original STACK question.
No external wording or mathematical
content is copied.
&lt;/p&gt;

&lt;p id="licence-dependency"&gt;
None.
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained outside
the question source.
&lt;/p&gt;
      </text>
    </questiondescription>


    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>


    <prtcorrect format="html">
      <text>&lt;p&gt;
Correct. A rounded result represents an
interval of possible original values.
&lt;/p&gt;</text>
    </prtcorrect>


    <prtpartiallycorrect format="html">
      <text>&lt;p&gt;
One boundary is correct. Recheck the
halfway points around the rounded value.
&lt;/p&gt;</text>
    </prtpartiallycorrect>


    <prtincorrect format="html">
      <text>&lt;p&gt;
A rounded value can represent many
different original numbers.
&lt;/p&gt;</text>
    </prtincorrect>


    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />


    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>lower_answer</tans>
      <boxsize>14</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>


    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>upper_answer</tans>
      <boxsize>14</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>


    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>

      <feedbackvariables>
        <text />
      </feedbackvariables>


      <node>
        <name>0</name>
        <description>
Checks the lower boundary.
        </description>

        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>lower_answer</tans>
        <testoptions />
        <quiet>0</quiet>

        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty />
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>

        <truefeedback format="html">
          <text />
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>

        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>


      <node>
        <name>1</name>
        <description>
Lower boundary is correct.
Checks the upper boundary.
        </description>

        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>upper_answer</tans>
        <testoptions />
        <quiet>0</quiet>

        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
Correct. You identified the complete
whole-number interval.
&lt;/p&gt;</text>
        </truefeedback>

        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>

        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>


      <node>
        <name>2</name>
        <description>
Checks whether the learner assumes the
rounded value is the only possible
original value.
        </description>

        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_unique_lower</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>

        <truefeedback format="html">
          <text />
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>

        <falsefeedback format="html">
          <text>&lt;p&gt;
Find the halfway point between this rounded
value and the previous multiple.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>


      <node>
        <name>3</name>
        <description>
Checks whether the upper halfway boundary
was incorrectly included.
        </description>

        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrong_upper_inclusive</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
You included the upper halfway value.
&lt;/p&gt;

&lt;p&gt;
Under the half-up rule, that exact halfway
value rounds to the next higher multiple,
so the largest whole number in this
interval is one less.
&lt;/p&gt;
          </text>
        </truefeedback>

        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>

        <falsefeedback format="html">
          <text>&lt;p&gt;
Recheck the upper halfway boundary.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>


      <node>
        <name>4</name>
        <description>
Checks whether both boundaries were set
equal to the rounded result.
        </description>

        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrong_unique_upper</tans>
        <testoptions />
        <quiet>1</quiet>

        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>

        <truefeedback format="html">
          <text>&lt;p&gt;
A rounded result does not identify one
unique original value.
&lt;/p&gt;

&lt;p&gt;
There is a whole interval of values that
round to the same result.
&lt;/p&gt;
          </text>
        </truefeedback>

        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>

        <falsefeedback format="html">
          <text>&lt;p&gt;
Find both boundary values around the
rounded target.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>

    </prt>


    <qtest>
      <testcase>1</testcase>

      <description>
Correct lower and upper boundaries receive
full credit.
      </description>

      <testinput>
        <name>ans1</name>
        <value>lower_answer</value>
      </testinput>

      <testinput>
        <name>ans2</name>
        <value>upper_answer</value>
      </testinput>

      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>
prt1-1-T | prt1-2-T
        </expectedanswernote>
      </expected>
    </qtest>


    <qtest>
      <testcase>2</testcase>

      <description>
Assuming the rounded result is the only
possible original value receives zero
credit.
      </description>

      <testinput>
        <name>ans1</name>
        <value>wrong_unique_lower</value>
      </testinput>

      <testinput>
        <name>ans2</name>
        <value>wrong_unique_upper</value>
      </testinput>

      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>
prt1-1-F | prt1-3-T | prt1-5-T
        </expectedanswernote>
      </expected>
    </qtest>


    <qtest>
      <testcase>3</testcase>

      <description>
Correct lower boundary with an inclusive
upper halfway point receives half credit.
      </description>

      <testinput>
        <name>ans1</name>
        <value>lower_answer</value>
      </testinput>

      <testinput>
        <name>ans2</name>
        <value>wrong_upper_inclusive</value>
      </testinput>

      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>
prt1-1-T | prt1-2-F | prt1-4-T
        </expectedanswernote>
      </expected>
    </qtest>


  <deployedseed>1638009549</deployedseed><deployedseed>805437244</deployedseed><deployedseed>148365353</deployedseed><deployedseed>613188564</deployedseed><deployedseed>1097990906</deployedseed></question>
</quiz>