<?xml version='1.0' encoding='UTF-8'?>
<quiz>
  <question type="stack">
    <name>
      <text>
Shared: Round a whole number to a specified place value
      </text>
    </name>
    <questiontext format="html">
      <text>&lt;p&gt;
Round the number
\[
{@number@}
\]
to the nearest
&lt;strong&gt;{@place_name@}&lt;/strong&gt;.
&lt;/p&gt;

&lt;p&gt;
[[input:ans1]]
[[validation:ans1]]
[[feedback:prt1]]
&lt;/p&gt;
      </text>
    </questiontext>
    <generalfeedback format="html">
      <text>&lt;h3&gt;Worked solution&lt;/h3&gt;

&lt;p&gt;
We are rounding
\[
{@number@}
\]
to the nearest
&lt;strong&gt;{@place_name@}&lt;/strong&gt;.
&lt;/p&gt;

&lt;p&gt;
The rounding place has value
\[
{@place_value@}.
\]
Look at the digit immediately to its right.
&lt;/p&gt;

&lt;p&gt;
The correct rounded value is
\[
\boxed{{@answer@}}.
\]
&lt;/p&gt;
      </text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.1000000</penalty>
    <hidden>0</hidden>
    <idnumber />
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>
/*
Shared Question Bank
Place Value, Rounding and Standard Form
PVRS-02

Foundation task:
round a whole number to a specified
place value.

This question is curriculum-independent.
*/


/*
Choose one of three rounding places:
1 = nearest ten
2 = nearest hundred
3 = nearest thousand
*/

place_choice:
    rand_with_prohib(1,3,[]);


/*
Resolve place value.
*/

place_value:
    if place_choice=1 then 10
    elseif place_choice=2 then 100
    else 1000;


/*
Generate a six-digit candidate.

Then ensure the digit immediately to the
right of the rounding place is at least 5.

This guarantees that correct rounding
requires rounding upward, so truncation
is always a distinguishable misconception.
*/

raw_number:
    rand_with_prohib(
        100000,
        998999,
        []
    );

remainder:
    mod(
        raw_number,
        place_value
    );

number:
    if remainder &lt; place_value/2
    then raw_number
        + place_value/2
        - remainder
    else raw_number;


/*
Display label.
*/

place_name:
    if place_choice=1 then "ten"
    elseif place_choice=2 then "hundred"
    else "thousand";


/*
Teacher answer.

Use explicit half-up rounding so a digit
of 5 always rounds upward, matching the
school rounding convention.
*/

answer:
    floor(
        (
            number
            + place_value/2
        )
        / place_value
    )
    * place_value;


/*
Target misconception:
learner truncates instead of rounding.
*/

wrong_truncate:
    floor(number/place_value)
    *place_value;
      </text>
    </questionvariables>
    <specificfeedback format="html">
      <text />
    </specificfeedback>
    <questionnote format="html">
      <text>&lt;p&gt;
Shared Question Bank |
Place Value and Rounding |
PVRS-02
&lt;/p&gt;

&lt;p&gt;
number={@number@}
&lt;/p&gt;

&lt;p&gt;
place_value={@place_value@}
&lt;/p&gt;

&lt;p&gt;
answer={@answer@}
&lt;/p&gt;
      </text>
    </questionnote>
    <questiondescription format="html">
      <text>&lt;p id="bank"&gt;
Shared Curriculum-Independent Question Bank
&lt;/p&gt;

&lt;p id="shared-question-id"&gt;
PVRS-02
&lt;/p&gt;

&lt;p id="knowledge-concept"&gt;
number-sense
&lt;/p&gt;

&lt;p id="secondary-concept"&gt;
whole-number-rounding
&lt;/p&gt;

&lt;p id="role"&gt;
foundation
&lt;/p&gt;

&lt;p id="interaction"&gt;
numeric-rounding
&lt;/p&gt;

&lt;p id="description"&gt;
Learners round a multi-digit whole number
to the nearest ten, hundred, or thousand.
&lt;/p&gt;

&lt;p id="misconception"&gt;
Distinguishes correct rounding from
simple truncation at the requested place.
&lt;/p&gt;

&lt;p id="provenance"&gt;
Original STACK question. Interaction design
was informed by open rounding-question
patterns. No source wording or mathematical
content is copied.
&lt;/p&gt;

&lt;p id="curriculum-policy"&gt;
Curriculum alignment is maintained outside
the question source.
&lt;/p&gt;
      </text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>&lt;p&gt;Correct.&lt;/p&gt;</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text />
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>&lt;p&gt;
Check the digit immediately to the right
of the place you are rounding to.
&lt;/p&gt;</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed />
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>answer</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint />
      <syntaxattribute>0</syntaxattribute>
      <forbidwords />
      <allowwords />
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options />
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text />
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>
Checks the correctly rounded value.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>answer</tans>
        <testoptions />
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text />
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text />
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>
Checks whether the learner truncated
instead of rounding.
        </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_truncate</tans>
        <testoptions />
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty />
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>&lt;p&gt;
It looks like you removed the digits after
the requested place instead of rounding.
&lt;/p&gt;

&lt;p&gt;
Look at the next digit:
if it is \(5\) or greater, increase the
rounding digit by \(1\).
&lt;/p&gt;
          </text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty />
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>&lt;p&gt;
Identify the requested place, then inspect
the digit immediately to its right.
&lt;/p&gt;</text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>
Teacher answer should receive full credit.
      </description>
      <testinput>
        <name>ans1</name>
        <value>answer</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>
prt1-1-T
        </expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Truncation misconception should receive zero credit with targeted feedback.</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_truncate</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F | prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>